# Optimization Exercise 1

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import scipy.optimize as opt
```

## Hat potential

The following potential is often used in Physics and other fields to describe symmetry breaking and is often known as the "hat potential":

$$ V(x) = -a x^2 + b x^4 $$

Write a function `hat(x,a,b)` that returns the value of this function:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: optimizationex01a, points: 2, grade: true, solution: false -->
```python
assert hat(0.0, 1.0, 1.0)==0.0
assert hat(0.0, 1.0, 1.0)==0.0
assert hat(1.0, 10.0, 1.0)==-9.0
```

Plot this function over the range $x\in\left[-3,3\right]$ with $b=1.0$ and $a=5.0$:

```python
a = 5.0
b = 1.0
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: optimizationex01b, points: 2, grade: true, solution: false -->
```python
assert True # leave this to grade the plot
```

Write code that finds the two local minima of this function for $b=1.0$ and $a=5.0$.

* Use `scipy.optimize.minimize` to find the minima. You will have to think carefully about how to get this function to find both minima.
* Print the x values of the minima.
* Plot the function as a blue line.
* On the same axes, show the minima as red circles.
* Customize your visualization to make it beatiful and effective.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: optimizationex01c, points: 4, grade: true, solution: false -->
```python
assert True # leave this for grading the plot
```

To check your numerical results, find the locations of the minima analytically. Show and describe the steps in your derivation using LaTeX equations. Evaluate the location of the minima using the above parameters.

<!-- grade_id: optimizationex01d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE