# Interpolation Exercise 1

```python
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
import numpy as np
```

```python
from scipy.interpolate import interp1d
```

## 2D trajectory interpolation

The file `trajectory.npz` contains 3 Numpy arrays that describe a 2d trajectory of a particle as a function of time:

* `t` which has discrete values of time `t[i]`.
* `x` which has values of the x position at those times: `x[i] = x(t[i])`.
* `x` which has values of the y position at those times: `y[i] = y(t[i])`.

Load those arrays into this notebook and save them as variables `x`, `y` and `t`:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interpolationex01a, points: 4, grade: true, solution: false -->
```python
assert isinstance(x, np.ndarray) and len(x)==40
assert isinstance(y, np.ndarray) and len(y)==40
assert isinstance(t, np.ndarray) and len(t)==40
```

Use these arrays to create interpolated functions $x(t)$ and $y(t)$. Then use those functions to create the following arrays:

* `newt` which has 200 points between $\{t_{min},t_{max}\}$.
* `newx` which has the interpolated values of $x(t)$ at those times.
* `newy` which has the interpolated values of $y(t)$ at those times.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interpolationex01b, points: 4, grade: true, solution: false -->
```python
assert newt[0]==t.min()
assert newt[-1]==t.max()
assert len(newt)==200
assert len(newx)==200
assert len(newy)==200
```

Make a parametric plot of $\{x(t),y(t)\}$ that shows the interpolated values and the original points:

* For the interpolated points, use a solid line.
* For the original points, use circles of a different color and no line.
* Customize you plot to make it effective and beautiful.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interpolationex01c, points: 2, grade: true, solution: false -->
```python
assert True # leave this to grade the trajectory plot
```