# Project Euler: Problem 17

https://projecteuler.net/problem=17

If the numbers 1 to 5 are written out in words: one, two, three, four, five, then there are 3 + 3 + 5 + 4 + 4 = 19 letters used in total.

If all the numbers from 1 to 1000 (one thousand) inclusive were written out in words, how many letters would be used?


NOTE: Do not count spaces or hyphens. For example, 342 (three hundred and forty-two) contains 23 letters and 115 (one hundred and fifteen) contains 20 letters. The use of "and" when writing out numbers is in compliance with British usage.

First write a `number_to_words(n)` function that takes an integer `n` between 1 and 1000 inclusive and returns a list of words for the number as described above

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def number_to_words(n):
    if(not(type(n)) == int): raise TypeError("number_to_words was passed a non-integer: " + str(n))
    if(n > 1000 or n < 1): raise Error("number_to_words was passed a value out of bounds [1,1000]: " + n)
        
    """Given a number n between 1-1000 inclusive return a list of words for the number."""
    units = ['one', 'two', 'three', 'four', 'five', 'six',  'seven', 'eight', 'nine']
    strbuild = ""
    if(n == 1000): return "onethousand"
    if(n >= 100):
        h = (n - (n % 100)) / 100 #Get hundreds place out of number
        n = n - 100 * h
        strbuild += units[int(h-1)] + "hundred"
        if(not(n == 0)): strbuild += "and";
    if(n >= 20):
        tens = ['ten', 'twenty', 'thirty', 'forty', 'fifty', 'sixty', 'seventy', 'eighty', 'ninety']
        t = (n - (n % 10)) / 10
        n = n - 10 * t
        strbuild += tens[int(t-1)]
    if(n >= 10):
        ltens = ['ten', 'eleven', 'twelve', 'thirteen', 'fourteen', 'fifteen', 'sixteen', 'seventeen', 'eighteen', 'nineteen']
        strbuild += ltens[int(n-10)]
        n = 0
    if(n >= 1):
        strbuild += units[int(n-1)]
    return strbuild.strip()
```

Now write a set of `assert` tests for your `number_to_words` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
assert type(number_to_words(190)) == str

# Available confirmations from the website
assert len(number_to_words(342)) == 23
assert len(number_to_words(115)) == 20
assert len(number_to_words(467)) == 24

assert number_to_words(100) == "onehundred"
assert number_to_words(102) == "onehundredandtwo"

# NOTE:
# Without pytest library (not included in Jupyter), 
# unable to test contract-based design
```

<!-- grade_id: projecteuler17a, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the number_to_words tests.
```

Now define a `count_letters(n)` that returns the number of letters used to write out the words for all of the the numbers `1` to `n` inclusive.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def count_letters(n):
    s = 0
    for i in range(1, n+1):
        s += len(number_to_words(i))
    return s
```

Now write a set of `assert` tests for your `count_letters` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
assert type(count_letters(1)) == int
assert count_letters(1) == 3
assert count_letters(2) == 6
assert count_letters(3) == 11
```

<!-- grade_id: projecteuler17b, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the count_letters tests.
```

Finally used your `count_letters` function to solve the original question.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
print(count_letters(1000))
```

Output:
```
21124
```

<!-- grade_id: projecteuler17c, points: 2, grade: true, solution: false -->
```python
assert True # use this for gradig the answer to the original question.
```